Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

# Week7 Assignment - Regression.

For this assignment, you will use the Ames housing dataset. The basic goal of this assignment is to predict housing prices for houses in Ames, Iowa.

In the first component, you will load the training dataset and convert it into a Pandas dataframe.

In the second component, you will plot the data. This component is ungraded.

In the third component, you will preprocess the data.

In the fourth component, you will split your data for training and testing.

In the fifth component, you will build and evaluate your model on the test data.

All the exercises are designed so that the solutions will need only one or a few lines of code.

Do not hesitate to contact instuctors and TA via #week7 channel on Slack if you get stuck. Join the channel first by clicking on Channels.

## Part A. Read in the Ames dataset. 

In this component you will read the Ames housing dataset into a Pandas data frame. This dataset contains samples of Ames houses, with multiple features. The value we want to predict is the cost of given house.

Make sure you download the ames_train.csv file from the course website. Do not manually modify the ames datasets.

Save to a data frame named ames_data.

<!-- grade_id: cell-c1e7f79ad332b8f9, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np

# Load the ames dataset into a pandas dataframe
# Make sure to save as "ames_data"

# YOUR CODE HERE
pass

# You can make sure you loaded it correctly by looking at its first few rows, using the .head() function
#print(ames_data.head())
```

<!-- grade_id: cell-fd50146eeead0c65, points: 1, grade: true, solution: false -->
```python
assert isinstance(ames_data, pd.core.frame.DataFrame)
assert ames_data.shape == (1460, 81)
assert ames_data["1stFlrSF"].sum() == 1697435
assert ames_data["2ndFlrSF"].sum() == 506609
assert ames_data.isnull().sum().sum() == 6965
```

## Part B. Plot the data.
Now that we've loaded the data, let's take a look at some features.

Perhaps naively, you think that the best predictors for house prices are the square footage values ("1stFlrSF", "2ndFlrSF").

Go ahead and plot a histogram of each, to get an idea of their distributions. See if you see anything interesting.

Following this, maybe try to get feel of the relationship between each of these and our target ("SalePrice", the house prices). Plot a scatter plot with SalePrice on the Y axis and 1stFlrSF on the X axis. Does it look like it may be a good predictor? What about 2ndFlrSF?

What about the sum of the two? Plot a scatter plot with SalePrice on the Y axis and the sum of 1stFlrSF and 2ndFlrSF on the X axis.

```python
import matplotlib.pyplot as plt
%matplotlib inline

```

## Part C. Preprocess the data.

It looks like the sum of 1stFlrSF and 2ndFlrSF may be good predictors.

Create a new column, named sqft_sum, that contains the summation of these two.

Look at the scatter plot of sqft_sum on the X axis and SalePrice on the Y axis. Note that it appears that there are two outliers: houses that are large but sold for very cheap. It may be a good choice to remove these, as it appears that they may have sold for much less than they are worth.

**Important note**: Whether or not to drop outliers is a debatable topic. Namely, it depends what you are trying to do and what the outliers represent. Wrongly treating outliers can bias your data. If you know the outliers are in error, or you know that they are not important to your goals, you can drop them. However, sometimes you will want to examine why they are outliers, what they are telling you. In our example, you may want to figure out WHY these houses sold for less than they are worth. E.g. Are they in a particularly loud neighborhood? Once we know that we can decide whether or not to include them. As stated in class, preprocessing is a very subjective phase. For the homework assignment, however, we will drop both of these outliers.

To summarize, in the next cell you need to:
1. Create a new column, sqft_sum, that is equal to the sum of the columns 1stFlrSF and 2ndFlrSF.
2. Remove the two clear outliers. If you plot the data (sqft_sum on X axis, SalePrice on Y axis), you should see them clearly in the plot. They are the two houses that appear to have been sold for much less than other houses in the same size range.

<!-- grade_id: cell-8d46731c6bc218b6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
pass

# Once you've altered ames_data, you can go ahead and replot the scatter plot to see if you indeed removed the data
# points you intended to remove.
```

<!-- grade_id: cell-c7edebc7b2cd1c7b, points: 1, grade: true, solution: false -->
```python
assert ames_data.shape == (1458, 82)
assert "sqft_sum" in ames_data.columns
assert ames_data["sqft_sum"].sum() == 2193726
```

## Part D. Split our data for evaluation.

We've loaded our data, plotted it to see what it looks like, and preprocessed it.

It's time to train and evaluate our model. At this stage, we will be evaluating based on a test set we create by splitting the data in two.

The features you will be using are: sqft_sum and Fireplaces (the number of fireplaces in the house).

Below, you will set X to equal the ames_data with only the two columns we are interested in (sqft_sum and Fireplaces), and y to equal the ames_data SalePrice column (what we are trying to predict). Then, you will run the code below to split the data into training and test.

<!-- grade_id: cell-377e7ae405383f99, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split

# YOUR CODE HERE
pass

# Now we will split your data into training and test, for evaluation
# Note the function below
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size = 0.25, random_state=12323)
```

<!-- grade_id: cell-fd44161086b39a45, points: 1, grade: true, solution: false -->
```python
assert X_train.shape == (1093, 2)
assert X_test.shape == (365, 2)
assert y_train.sum() == 196281710
```

## Part E. Build and evaluate a linear regression model.
Build a linear regression model based on the X_train data. Name the model *lr*.

Once the model is built, predict the SalePrice for the test set X_test. Name the predictions *y_preds*.

Calculate the RMSE between y_preds and y_test. Name the RMSE *rmse*.


Hints:
* Use sklearn's LinearRegression function to fit a model.
* Use sklearn's mean_squared_error to evaluate the model.

<!-- grade_id: cell-b3bd68da8483ec11, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LinearRegression
from sklearn.metrics import mean_squared_error

# YOUR CODE HERE
pass
```

<!-- grade_id: cell-346d5411d0d8d874, points: 1, grade: true, solution: false -->
```python
assert isinstance(lr, LinearRegression)
assert np.isclose(y_preds.sum(), 67114985.983310223)
assert np.isclose(rmse, 53444.9760401)
```