# Week1 Assignment - Testing OKPy and the autograder.

For this assignment, you will load the iris flowers dataset from sklearn and predict flower species based on their petal and sepal sizes.

## Part A. Load the Iris dataset.

Sklearn comes with prepared datasets that you can load. To load the iris dataset, run:
```python
from sklearn import datasets
iris = datasets.load_iris()
```

<!-- grade_id: cell-c1e7f79ad332b8f9, points: 0, grade: false, solution: true -->
```python
from sklearn import datasets

# Load the iris dataset
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fd50146eeead0c65, points: 1, grade: true, solution: false -->
```python
assert (iris['data'] == datasets.load_iris()['data']).all()
```

## Part B. Use an SVM to predict iris species.

To setup the data and train an SVM, run the following:
```python
from sklearn.svm import SVC

X = iris.data
y = iris.target

classifier = SVC(kernel="rbf")
classifier.fit(X, y)
predictions = classifier.predict(X)
```

<!-- grade_id: cell-bc8ce57a60e5fd7a, points: 0, grade: false, solution: true -->
```python
from sklearn.svm import SVC

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-165c467d2a6e7dde, points: 1, grade: true, solution: false -->
```python
assert X.shape == (150, 4)
assert len(y) == 150
```