Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

# Week 10 - Assignments
## Dimensionality reduction and feature selection

There are two assignements based on examples shown in the classroom. Each completed assignment gives one point. There are also two ungraded optional assignments: one applies NMF to faces dataset and the second one applies clustering to the BRCA dataset analyzed in the classroom.



```python
import numpy as np
import matplotlib.pyplot as plt
%matplotlib inline
```

## 1. Task 1 (1 point)
Apply feature selection to the [Olivetti faces dataset](http://scikit-learn.org/stable/modules/generated/sklearn.datasets.fetch_olivetti_faces.html#sklearn.datasets.fetch_olivetti_faces), identifying the most important 25% of features based on the chi2 criteria (you can use SelectPercentile). This dataset consists of 10 pictures each of 40 individuals. Each picture is 64x64 pixels, thus the dimensionality of the problem is 4096 features if we consider that each pixel is a feature . You can visualize these top selected features on a 64x64 pixel heatmap for visual aid.

What the test expects:
 * Use variables __X__ and __y__ for data and the corresponding labels.
 * The new dataset with 25% of features should be named __X_new__ 
    

```python
from sklearn.datasets import fetch_olivetti_faces
from sklearn.feature_selection import SelectPercentile, chi2
```

<!-- grade_id: cell-38048a5ad8a510c4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
pass
```

<!-- grade_id: cell-2b447393aad4f719, points: 1, grade: true, solution: false -->
```python
assert X.shape == (400, 4096)
assert y.shape == (400,)
assert X_new.shape == (400, 1024)

assert min(chi2(X_new, y)[0]) > min(chi2(X, y)[0])

####################################### 
# chi2 score of each feature (pixel)
importances = np.array(chi2(X, y)[0])
plt.matshow(importances.reshape(64,64), cmap=plt.cm.hot)

# only top 25% of the most important features
min_chi2 = min(chi2(X_new, y)[0])
importances = np.where(importances < min_chi2, np.inf, importances)
plt.matshow(importances.reshape(64,64), cmap=plt.cm.hot)
```

### Optional exercise 1 (not graded) - NMF applied to faces

Have a look at the seminal paper that introduced NMF: http://www.columbia.edu/~jwp2128/Teaching/E4903/papers/nmf_nature.pdf

I hope you will become inspired to reproduce the original experiment by applying NMF to Olivetti faces dataset. Decompose faces into a desired number of components (try 8, then 16, for instance). Show the components as images. Due to non-negative nature of the procedure, the individual components can be interpreted as facial features.


<!-- grade_id: cell-8f082768410bca19, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
pass
```

## Task 2 (1 point)

Apply tSNE to the digits dataset (transform to 2D). Digits dataset has about 180 samples per class, so we should set perplexity=150 in tSNE accordingly. Afterwards, perform K-means clustering into 10 clusters in 2D feature space obtained after tSNE transformation. Calculate sklearn.metrics.adjusted_rand_score between the labels obtained in clustering and the original digit labels. http://scikit-learn.org/stable/modules/generated/sklearn.datasets.load_digits.html

What the test expects:
 * Use variables __X__ and __y__ for data and the corresponding labels.
 * tSNE-transformed 2D dataset should be named __X_new__
 * The variable holding Adjusted Rand Score should be named __score__
 * The variable holding cluster labels assigned by KMeans should be named __y_new__

```python
from sklearn.datasets import load_digits
from sklearn.manifold import TSNE
from sklearn.cluster import KMeans
from sklearn.metrics import adjusted_rand_score
```

<!-- grade_id: cell-ecffeb3f1e71f096, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
pass
```

<!-- grade_id: cell-f5290f330ad9a78d, points: 0, grade: true, solution: false -->
```python
assert X.shape == (1797, 64)
assert X_new.shape == (1797, 2)
assert y.shape == (1797, )
assert y_new.shape == (1797, )
print(score)
assert score > 0.7

###################################
import seaborn as sns
import pandas as pd

df = pd.DataFrame({'tSNE 1': X_new[:, 0], 'tSNE 2': X_new[:, 1], 'actual_label': y, 'cluster': y_new})
sns.lmplot(data=df, x='tSNE 1', y='tSNE 2', hue='actual_label', fit_reg=False)
sns.lmplot(data=df, x='tSNE 1', y='tSNE 2', hue='cluster', fit_reg=False)
```

## 3. Optional exercise 2 (not graded) - Comparing clustering before and after NMF

Use BRCA gene expression data (measured with antibodies on a protein level)  provided in this class. Apply K-means clustering to the original expression data first; then apply it to the result of NMF decomposition with 10 components (H matrix). For simplicity, also use k=10 as the number of clusters in K-means. Calculate adjusted_rand_score as a measure of consensus between the two clustering results.



<!-- grade_id: cell-6dcbc41d73f5f97c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
pass
```