Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

# Week 2 Assignments
The assignment for week 2 has four components, each worth 25%. Each component is graded as pass or fail. 

In the first, you will implement a method for a class that represents a rectangle. The method you implement will return the rectangle's size.

In the second component, you will create two classes, "Dog" and "Cat", that inherit from the abstract "Pet" class that we provide. You will implement a method for each of these classes that returns a string describing the noise the particular pet makes. Namely, "bark" for dog and "meow" for cat.

The third component will require that you create a class from scratch, that represents a bank account. You will need to implement methods to deposit, withdraw and check the account balance.

In the fourth and final component, you will create a class "DebitAccount" that is identical to BankAccount, except it cannot accept a negative balance.

Below each assignment is a cell with several assertions. Run that cell to verify your code. If it fails, correct your code and rerun the cell containing the assertions. Once the cell containing the assertions runs correctly, your code is working. Note however that we will be running a few other tests to verify your code.

## Part A: Creating a Class (1 point)

Write a class that represents a rectangle. We've written the first part for you. The rectangle's \_\_init\_\_ function
takes width and height parameters, so:
```python
my_rectangle = Rectangle(2, 4)
```
creates a 2 x 4 rectangle.

The class needs a method that returns the rectangle's size. The method: 
* Takes no arguments (other than "self")
* Returns the rectangle's size (height * width)

IE running the following code will give you the size:
```python
my_rectangle = Rectangle(2, 4)
print(my_rectangle.get_size())
```
Should print: 8.

<!-- grade_id: p1, points: 0, grade: false, solution: true -->
```python
class Rectangle:
    def __init__(self, height, width):
        """Initialize a new rectangle object."""
        self.height = height
        self.width = width
        
    def get_size(self):
        """Return the size of the rectangle."""
        # YOUR CODE HERE
        raise NotImplementedError()
```

Your rectangle should indicate a size of 20 for height of 2 and width of 10. Check that it does:

```python
my_rectangle = Rectangle(2, 10)
assert my_rectangle.get_size() == 20
```

The next set of assertions should run without error. Run them to verify your code.

<!-- grade_id: correct_p1, points: 1, grade: true, solution: false -->
```python
assert Rectangle(10, 10).get_size() == 100
assert Rectangle(10, 30).get_size() == 300
assert Rectangle(28, 2).get_size() == 56
assert Rectangle(10, 10).get_size() == 100
```

---
## Part B: Abstraction, inheritance and polymorphism (1 point)

Below we've created an abstract class to represent pets. Its \_\_init\_\_ function takes a name for the pet.

You will inherit from the "Pet" class to create a "Dog" and a "Cat". You will override the "make_sound" method, so that for a dog it returns the string "bark" and for a cat it returns the string "meow".

IE running the following code will give you the expected sounds:
```python
dog = Dog("sherlock")
print(dog.make_sound())
```
The above code should print: "bark".

<!-- grade_id: p2, points: 0, grade: false, solution: true -->
```python
class Pet:
    def __init__(self, name):
        """
        Initialize a new pet object.
        The name parameter is the pet's name.
        """
        self.name = name
    
    def make_noise(self):
        """Method to make noise. Should be overridden by the inheriting class."""
        pass

class Dog(Pet):
    # YOUR CODE HERE
    raise NotImplementedError()

class Cat(Pet):
    # YOUR CODE HERE
    raise NotImplementedError()
```

The next set of assertions should run without error. Run them to verify your code.

<!-- grade_id: correct_p2, points: 1, grade: true, solution: false -->
```python
pet1 = Dog("sherlock")
assert pet1.make_noise() == "bark"
assert pet1.name == "sherlock"

pet2 = Cat("rex")
assert pet2.make_noise() == "meow"
assert pet2.name == "rex"
```

---
## Part C:  Object state (1 point)

You will need to create a class "BankAccount". Its \_\_init\_\_ function should initialize the attribute that will hold the current balance amount.

You will need to implement methods "deposit", "withdraw" that take amount as an argument (a non-negative float) and adjust the account balance accordingly. You will also need to implement method "check_balance" that returns current account balance.

Make sure to create appropriate docstring documentation for each method and also for the class. The docstring documentation is placed immediately below the class declaration and below each method's declaration, with three quotation marks on each side, such that:
```python
class MyClass:
    """
    This is where documentation on what
    MyClass does should go.
    """
    
    def my_method(self):
        """
        This is where documentation on what 
        my_method does should go.
        """
        pass
```

<!-- grade_id: p3, points: 0, grade: false, solution: true -->
```python
class BankAccount:
    # YOUR CODE HERE
    raise NotImplementedError()
```

The next set of assertions should run without error. Run them to verify your code.

<!-- grade_id: correct_p3, points: 1, grade: true, solution: false -->
```python
acc = BankAccount()
acc.deposit(100.0)
acc.deposit(200.0)
acc.withdraw(250.0)
assert acc.check_balance() == 50.0
acc.withdraw(10.0)
acc.deposit(30.0)
assert acc.check_balance() == 70.0
assert len(BankAccount.deposit.__doc__) > 10
assert len(BankAccount.withdraw.__doc__) > 10
```

---
## Part D:  Inheritance and Polymorphisms (1 point)

You will need to create a class "DebitAccount" that behaves exactly like "BankAccount" with one notable exception -- it can not have a negative balance.

Therefore, the method "withdraw" of "DebitAccount" should not change the account balance if there are insufficient funds and should return False in this case.

If there are sufficient funds "withdraw" should adjust the account balance and return True.

Apply the concepts of inheritance and polymorphisms here.

As always, make sure to provide a docstring documentation, it is a requirement.

<!-- grade_id: pd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

The next set of assertions should run without error. Run them to verify your code.

<!-- grade_id: correct_pd, points: 1, grade: true, solution: false -->
```python
assert issubclass(DebitAccount, BankAccount)

debit = DebitAccount()
debit.deposit(30.0)
debit.deposit(20.0)
assert debit.check_balance() == 50.0
result = debit.withdraw(10.0)
assert debit.check_balance() == 40.0
assert result is True

result = debit.withdraw(90.0)
assert debit.check_balance() == 40.0
assert result is False

```

```python

```