# Week 3 Assignment

We encourage you to ask questions about Week 3 topic and this assignment on Slack channel #week3. If you have not joined the channel, click on "Channels" and choose the channels to join. 

The assignment for __Week 3__ has 6 components. The first component is a self-test quiz. It is not graded and is provided for self checking. The next three exercises are each worth 1 point, and each of them is graded as pass or fail. In order to pass you need to implement the required functionality and pass the assert tests. If you think that some assertions are incorrect please contact us on Slack #week3, it sometimes happens. The last two exercises are not graded but we highly recommend attempting to solve them. Depending on your Python skills, they may require some documentation reading, but this is a skill you all need to develop. 

Make sure that there are no errors, i.e. exceptions raised anywhere in the notebook. Exception in one cell will stop notebook execution and the code in all the following cells will not run.

Official Python documentation and tutorials is your main reference. For instance, Python reference on Exceptions https://docs.python.org/3/library/exceptions.html and Python tutorial on Classes https://docs.python.org/3/tutorial/classes.html

There are several online texbooks that may be helpful in understanding the object oriented programming concepts and the practical aspects of OOP in Python http://python-textbok.readthedocs.io/en/1.0/Classes.html  and https://python-textbok.readthedocs.io/en/1.0/Object_Oriented_Programming.html

If you are interested in learning about how OOP approach is applied in Software Engineering and Design, there is a great introduction by IBM https://www.ibm.com/developerworks/library/ws-using-oo/

In this course we focus mainly on the aspects that will become important in applying data analysis and machine learning Python packages. We assume that you already have some basic knowledge of the Python language and are comfortable programming in it.

## Part A: self-test quiz (not graded)

### 0. Which of the following represents a distinctly identifiable entity in the real world?
- a. A class
- b. An object
- c. A method
- c. An attribute

### 1. An object is a data entity, and any program that needs to can access any of that object's data.
- a. True
- b. False
- c. Sometimes

### 2. Which statement most accurately describes the relationship between a parent and child classes?
- a. A parent class contains the exact same code as its child.
- b. A child class is unrelated to its parent.
- c. A parent class inherits attributes and behavior from its child.
- d. A child class contains code from its parent without copying it, is more specialized and can define additional attributes.

### 3. A well-written object class can be described as:
- a. Has well-defined boundaries, does only a few things, and knows only about its own data.
- b. Is largely data oriented and open to any program that needs to access it.
- c. Has knowledge only of the other objects and classes it needs to perform its functions.
- d. a and b
- e. a and c
- f. All of the above

### 4. Which statement(s) best describes an object's attributes?
- a. Attributes act like the nouns of an object.
- b. Attributes describe an object's behavior.
- c. Attributes can mean pretty much whatever you want them to mean — that's what we mean by "open source."
- d. a and c
- e. b and c
- f. All of the above.

### 5. Which statement(s) best describes an object's behavior (methods)?
- a. Behavior acts like the nouns of an object.
- b. Behavior, like attributes, can be defined however we want — that's what we mean by "open source."
- c. Behavior acts like the verbs of an object.
- d. All of the above.

### 6. Encapsulation means "open to all" and indicates that the internal workings of an object are entirely visible to the outside world.
- a. True
- b. False
- c. Sometimes

### 7. Which statement(s) best describes inheritance?
- a. One class (the child) generalizes the attributes and behavior of another class (the parent), making the child object more generalized.
- b. The child class inherits from its parent, adding specialized attributes and behavior.
- c. When a child class is created, a well-known best practice is for the developer to manually copy/paste all of the source code from the parent class to ensure that nothing gets broken.
- d. When a child class inherits from its parent, only the additional attributes and behavior must be coded, which is one of the benefits of OOP.
- e. a and b
- f. b and d
- g. All of the above.


### 8. What will be the output of the following code snippet?

    class Test:
        def __init__(self):
            self.x = 1
            self.__y = 1

        def getY(self):
            return self.__y

    val = Test()
    print(val.__y)

- a. The program has an error because y is private and should not access it from outside the class. 
- b. The program has an error because you cannot name a variable using \_\_y. 
- c. The program runs fine and prints 1. 
- d. The program runs fine and prints nothing.


#### Answers to quiz

0. b 1. b 2. d 3. e 4. a 5. c 6. b 7. f 8. a

## Part B: Creating and applying a custom exception, working with strings (1 point)

Define a new class __SSNException__ that inherits from __Exception__ class. We will use this class to raise exceptions in case of SSN errors. This class does not have to modify behavior or attributes of the parent __Exception__ class.
Create a function __verify_ssn__ that accepts SSN as a string argument and raises an __SSNException__ exception if the string does not match the "000-00-0000" pattern. If the SSN format is correct the function does nothing.

As always, provide docstring documentation for each class and function you define.

<!-- grade_id: cell-891a23fd09ddc413, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-d29509f3af30d4b0, points: 1, grade: true, solution: false -->
```python
assert callable(verify_ssn)
assert issubclass(SSNException, Exception)

try:
    _saved_exception = None
    verify_ssn("")
except Exception as _e:
    _saved_exception = _e
assert type(_saved_exception) is SSNException

try:
    _saved_exception = None
    verify_ssn("123-45-6789")
except Exception as _e:
    _saved_exception = _e

assert _saved_exception is None

try:
    _saved_exception = None
    verify_ssn("123-45-")
except Exception as _e:
    _saved_exception = _e

assert type(_saved_exception) is SSNException

try:
    _saved_exception = None
    verify_ssn("123-00-45AA")
except Exception as _e:
    _saved_exception = _e

assert type(_saved_exception) is SSNException
```

## Part C: Initializing objects, using static attributes and overriding magic methods (1 point)

Define a new class __Client__ that would describe a client's personal information as object instance attributes __name__ and __SSN__. These attributes should be different for each client, i.e. for each __Client__ instance. 

Example of usage:

    client = Client(name="John Johnson", ssn="111-22-3333")

While initializing the object for __Client__ you should call __verify_ssn__ function to validate the correctness of the SSN provided.

Client class should also have a class attribute __organization__ that is going to be the same for all __Client__ objects once it is set. Sometimes such attributes are also called static.

Do not forget to provide docstring documentation.

Optionally, make __Client__ object instance attributes hidden so that they could not be accessed directly (hint: self.__ssn).

Optionally, define a method __\_\_str\_\___ that would allow printing client information in user convenient way. So that print(client) would return something like "Microsoft client John Johnson (SSN xxx-xx-3333)", where "Microsoft" in our example is the name of organization. Note that __\_\_str\_\___ should return a string, not print it.

<!-- grade_id: cell-911bbab271d62bff, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE


# usage example
client = Client("John Johnson", "111-22-3333")
client.organization = "Microsoft"
print(client)

```

Output:
```
Microsoft client John Johnson (SSN xxx-xx-3333)
```

<!-- grade_id: cell-4b185e580de1a6b0, points: 1, grade: true, solution: false -->
```python
_client = Client("John Doe", "123-00-4444")
assert type(_client) is Client

try:
    _saved_exception = None
    Client("John Doe", "123-00-45AA")
except Exception as _e:
    _saved_exception = _e

assert type(_saved_exception) is SSNException

assert _client.organization == Client.organization
```

## Part D: Creating a Class (1 point)

For this exercise we will reuse the code from the Week 2 assignment, namely the classes __BankAccount__ and __DebitAccount__. You will find that they are already defined in this notebook.

Define a new class __BankClient__ that inherits from __Client__ class from the previous exercise and that would describe a client of a bank. Set the __organization__ class attribute of __BankClient__ to "Bank of Bethesda".

Let us extend the capabilities of this class in comparison to its parent. __BankClient__ objects should have a container for one or several accounts compatible with __BankAccount__ interface. We will need a method __add_accounts(self, accounts)__ to assign the bank accounts to a bank client object. The __accounts__ parameter should be a list with one or several __BankAccount__ objects. The software design pattern where an object contains references to other objects is called *composition*, and now you are familiar with it.

Additionally, define a method __check_balance__ that will sum up all the balances from the associated accounts and return the resulting value.

Optionally, you can redefine __\_\_str\_\___ method to include the client's balance.

Also note that you can now practice "duck typing" approach, because both bank account and bank client will have a __check_balance__ function. If it quacks like a duck ...

Do not forget about docstrings.

```python
# Classes from Week 2

class BankAccount:
    """ Class that handles bank account operations """

    def __init__(self):
        """ Initial balance is zero """
        self.balance = 0.0
    
    def deposit(self, amount):
        """ Deposit amount to account """
        self.balance += amount
    
    def withdraw(self, amount):
        """ Withdraws amount from account """
        self.balance -= amount
    
    def check_balance(self):
        """ Returns current balance on account """
        return self.balance


class DebitAccount(BankAccount):
    """ Class that handles debit account operations """

    def withdraw(self, amount):
        """ Withdraws amount from account only if there are sufficient funds """
        if self.balance - amount < 0.0:
            return False

        self.balance -= amount
        return True
```

<!-- grade_id: cell-dd3dadd69eca7f2d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE


# usage example:
john = BankClient("John Johnson", "111-22-3333")
print(john)

acc = BankAccount()
acc.deposit(100.0)
acc.deposit(200.0)
acc.withdraw(250.0)

debit = DebitAccount()
debit.deposit(30.0)
debit.deposit(20.0)
debit.withdraw(10.0)

john.add_accounts([acc, debit])
print(john)
```

Output:
```
Bank of Bethesda client John Johnson (SSN xxx-xx-3333) has 0 accounts with balance 0.0
Bank of Bethesda client John Johnson (SSN xxx-xx-3333) has 2 accounts with balance 90.0
```

<!-- grade_id: cell-2467cadab667f08b, points: 1, grade: true, solution: false -->
```python
assert issubclass(BankClient, Client)

_client = BankClient("John Doe", "123-00-4444")
assert type(_client) is BankClient

try:
    _saved_exception = None
    Client("John Doe", "123-00-45AA")
except Exception as _e:
    _saved_exception = _e

assert type(_saved_exception) is SSNException

assert _client.organization == "Bank of Bethesda"

assert callable(_client.add_accounts)

acc = BankAccount()
acc.deposit(100.0)
acc.deposit(200.0)
acc.withdraw(250.0)

debit = DebitAccount()
debit.deposit(30.0)
debit.deposit(20.0)
debit.withdraw(10.0)

_client.add_accounts([acc, debit])

assert _client.check_balance() == 90.0
```

## Part E:  DNA Sequence (not graded)

This is a creative assignment that is not graded and is designed to expose you to a more open-ended problem which uses object-oriented approach. In this exercise we will become more familiar with overriding the "magic" methods and attributes that correspond to Python operators and are part of the syntax of Python programming language. 

Create a class for DNA sequence that essentially is an extension on top of the standard Python string class. Let us give it a name DNASequence. We would like it to be an iterator (https://wiki.python.org/moin/Iterator) , i.e. it should implement __\_]_iter\_\_()__ and __next()__ methods.

We also want the new class to be reversible, i.e. implement __\_\_reversed\_\_()__ method. However, it should return a reverse complement DNA sequence (like https://www.bioinformatics.org/sms/rev_comp.html)
The __\_\_init\_\_()__ method should raise ValueError exception if the sequence contains anything else than A T G or C.


<!-- grade_id: cell-fa14022db83dd522, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE

```

```python
"""
# a possible test could be:
dna = DNASequence("ATGACACTCCGTGTTTAG")
complements = {'C': 'G', 'G':'C', 'A': 'T', 'T': 'A'}
for base, rc_base in zip(dna, reversed(dna)):
    assert base in complements.keys()
    assert base == complements[rc_base]
""";
```

## Part F: Implementing defaultdict (not graded)

This is the second creative assignment that is also not graded.

Create an implementation of an extension to Python dictionary data type that is equivalent to "defaultdict"  https://docs.python.org/3/library/collections.html#collections.defaultdict . Read the documentation and implement the same functionality in your code.

We will call the new class __NewDict__ and let it inherit from __dict__ class. You need to implement __\_\_missing\_\_(key)__ method and define __\_\_init\_\_()__ constructor that accepts a callable "factory" for initializing the default values for missing keys in the dictionary.

It is a very convenient extension to conventional Python __dict__ that does not require the values to be assegned beforehand, so that you can access values by keys right away as if they have already been initialized. See the examples below.

Have fun!

<!-- grade_id: cell-2141d87b6a457272, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE


```

```python
"""
# Possible use cases:

d = NewDict(int)
print(d['not existing'] + 5)
d['key'] += 4
d['key'] += 6
print(d['key'])

# An dictionary of lists
lst = NewDict(list)
lst['shopping_list'].append('milk')
lst['shopping_list'].append('bread')

# A counter for nucleotides
cnt = NewDict(int)
for i in "ATGACACTCCGTGTTTAG":
    cnt[i] += 1
print(cnt)

""";
```

```python

```