Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

# Week 8 - Latent Variable Models
For this week's assignment, you will implement a Gaussian Mixture Model and perform Factor Analysis

```python
from sklearn import cluster, datasets
import numpy as np
import pandas as pd
```

## Task 1. Load the Old Faithful dataset (1 pt)
For this assignment, we will be working with the Old Faithful dataset. Old Faithful is a geyser that is found in Yellowstone National Park, Wyoming and is the park's most famous attraction. We will be working with a dataset from 1990 that recorded the time between eruptions and the duration of the eruption, both taken in minutes. There are 272 observations in total.

We have provided the url to the dataset below. Use the pandas `read_csv()` function to read in the data. 

**NOTE:** you may want to use the `skipinitialspace` argument to properly load the header. 

<!-- grade_id: cell-9ce296dedde4b74f, points: 0, grade: false, solution: true -->
```python
# URL for the dataset
url = "https://raw.githubusercontent.com/barneygovan/from-data-with-love/master/data/faithful.csv"

# YOUR CODE HERE
pass
```

<!-- grade_id: cell-7ee94fa68dd0a648, points: 1, grade: true, solution: false -->
```python
assert data.shape == (272, 2)
```

## Task 2: Visualize the data (1 pt)
Let's get a sense of what our data look like by plotting the distribution of eruption times and waiting times, as well as a scatter plot of both dimensions.
<ol>
    <li>A histogram of the eruption time (the column labeled "eruptions").</li>
    <li>A histogram of the waiting time between eruptions.</li>
    <li>A scatter plot of both features</li>
</ol>

<!-- grade_id: cell-c9fcee0d9c5935b8, points: 1, grade: true, solution: true -->
```python
import matplotlib.pyplot as plt
%matplotlib inline
# YOUR CODE HERE
pass
```

## Task 3: Build a GMM (1 pt)
Each feature seems to be bimodally distributes and the scatterplot clearly shows separation of two clusters. 
<ol>
    <li> Using sklearn's `GaussianMixture` implementation, build a Gaussian Mixture Model with 2 Gaussians. Save the result in a variable named `gmm` <br>
**Note:** In order for our assertions to work, make sure you use the default parameters for `GaussianMixture`, with the exception of random_state, which you should set to 126
    </li>
    <li> Annotate each datapoint using the `.predict` function 
    <li> Then create a scatterplot of eruption time versus waiting time in which each point is colored based on its prediction from the GMM</li>
</ol>



<!-- grade_id: cell-1a6509d9470ae7e8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
pass
```

<!-- grade_id: cell-adc61c99588dcb78, points: 1, grade: true, solution: false -->
```python
assert labels.shape == (272,)
assert np.isclose(np.mean(labels), .64, .1)
```

## Task 4: Choosing N (ungraded)
As discussed in lecture, the Akaike information criterion (AIC) and the Bayesian information criterion (BIC) provide quantitative ways to choose the number of clusters that maximize the likelihood of our data while penalizing for increased complexity. Scikit-Learn's GMM estimator actually includes built-in methods that compute both of these.

Let's look at the AIC and BIC as a function of the number of GMM components for our dataset. Run the cell below to visualize the AIC and BIC metrics for 9 different models.  

```python
N = 21
n_gaussians = np.arange(1, N + 1)

# create a Gaussian mixture model for the range 1-10
models = [GaussianMixture(n, random_state=126).fit(data)
          for n in n_gaussians]

# plot the AIC and BIC:
plt.plot(n_gaussians, [m.bic(data) for m in models], label='BIC')
plt.plot(n_gaussians, [m.aic(data) for m in models], label='AIC')
plt.legend(loc='best')
plt.ylabel('AIC and BIC values');
plt.xlabel('N (number of GMM components)');
```

The optimal number of clusters is the value that minimizes the AIC or BIC. We can see from both of these measures that 2 seemed to be the right choice!

Now try these two experiments:
1. Change the value of N to 21 and rerun the code. Would we still prefer N=2 as an optimal value?
2. Change the value of N to 70 and rerun the analysis.

What do you think is happening and how should it be interpreted?

<!-- grade_id: cell-e2c429621294d825, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Task 5: Factor Analysis (1 pt)
Factor analysis is a useful technique for reducing dimensionality of data. It assumes that multiple observed variables have similar patterns of responses because they are all associated with a latent variable. 

In the 'beer dataset' provided, consumers (n = 99) rate on a scale of 0-100 how important they consider each of seven qualities when deciding whether or not to buy a 6-pack of beer, like cost, volume (size), alcohol percentage, the color, aroma, and taste. 

In future lectures, we will discuss how to go about choosing the right number of latent variables (factors, components) to use. However, for this assignment, we will perform factor analysis using 2 latent variables to see whether this data can be represented in 2 dimensions.

Steps:
1. Normalize the beer data using sklearn's `sklearn.preprocessing.scale` function </li>
2. Perform Factor Analysis with `sklearn.decomposition.FactorAnalysis` using 2 components. Save the resulting model using the variable name `fa`</li>
3. Observe the factor loadings, which you can access via the model's attribute `components_`. </li>
4. Visualize the factor loadings for each latent variable using a bar chart </li>

**Note:** In order for our assertions to work, make sure you use the default parameters for GaussianMixture with the exception of random_state that has to be set to 126.

<!-- grade_id: cell-0e5c27704d4bd093, points: 0, grade: false, solution: true -->
```python
url = "https://github.com/biof509/spring2019/raw/master/week8/beer.txt"
beer = pd.read_csv(url, sep="\t", dtype=np.float64)
print(beer.head(3))
print(beer.describe())

# YOUR CODE HERE
pass

```

<!-- grade_id: cell-4e7a6d78a0147e55, points: 1, grade: true, solution: false -->
```python
assert fa.components_.shape == (2, 6)
assert np.all(
        np.isclose(
            np.sum(fa.components_, axis=1),
            np.array([2.5, -3.1]),
            atol=.2))
```