Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

# Week4 Assignment - Regression.

For this assignment, you will use the Ames housing dataset. The basic goal of this assignment is to predict housing prices for houses in Ames, Iowa.

In the first component, you will load the training dataset and convert it into a Pandas dataframe.

In the second component, you will plot the data. This component is ungraded.

In the third component, you will preprocess the data.

In the fourth component, you will build a linear regression model.

In the fifth component, you will build a random forest model.

All the exercises are designed so that the solutions will need only one or a few lines of code.

Do not hesitate to contact instuctors and TA via #week4 channel on Slack if you get stuck. Join the channel first by clicking on Channels.

## Part A. Read in the Ames dataset. 

In this component you will read the Ames housing dataset into a Pandas data frame. This dataset contains samples of Ames houses, with multiple features. The value we want to predict is the cost of given house.

Make sure you download the ames_train.csv file from the course website. Do not manually modify the ames datasets.

Save to a data frame named ames_data.

<!-- grade_id: cell-c1e7f79ad332b8f9, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np

# Load the ames dataset into a pandas dataframe
# Make sure to save as "ames_data"

# YOUR CODE HERE
pass

# You can make sure you loaded it correctly by looking at its first few rows, using the .head() function
#print(ames_data.head())
```

<!-- grade_id: cell-fd50146eeead0c65, points: 1, grade: true, solution: false -->
```python
assert isinstance(ames_data, pd.core.frame.DataFrame)
assert ames_data.shape == (1460, 81)
assert ames_data["FirstFlrSF"].sum() == 1697435
assert ames_data["SecondFlrSF"].sum() == 506609
assert ames_data.isnull().sum().sum() == 6965
```

## Part B. Plot the data.
Now that we've loaded the data, let's take a look at some features.

Perhaps naively, you think that the best predictors for house prices are the square footage values ("FirstFlrSF", "SecondFlrSF").

Go ahead and plot a histogram of each, to get an idea of their distributions. See if you see anything interesting.

Following this, maybe try to get feel of the relationship between each of these and our target ("SalePrice", the house prices). Plot a scatter plot with SalePrice on the Y axis and FirststFlrSF on the X axis. Does it look like it may be a good predictor? What about SecondFlrSF?

What about the sum of the two? Plot a scatter plot with SalePrice on the Y axis and the sum of FirstFlrSF and SecondFlrSF on the X axis.

```python
import matplotlib.pyplot as plt
%matplotlib inline

```

## Part C. Preprocess the data.

It looks like the sum of FirstFlrSF and SecondFlrSF may be good predictors.

Create a new column, named sqft_sum, that contains the summation of these two.

<!-- grade_id: cell-8d46731c6bc218b6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
pass
```

<!-- grade_id: cell-c7edebc7b2cd1c7b, points: 1, grade: true, solution: false -->
```python
assert ames_data.shape == (1460, 82)
assert "sqft_sum" in ames_data.columns
assert ames_data["sqft_sum"].sum() == 2204044
```

## Part D. Build a linear regression model.
Build a linear regression model based on the Ames data. The features you will be using are: sqft_sum and Fireplaces (the number of fireplaces in the house).

Name the model *lr*.

Once the model is trained, predict the SalePrice for the same data you trained on. Name the predictions *y_preds*.

Hints:
* Use sklearn's LinearRegression function to fit a model.

<!-- grade_id: cell-b3bd68da8483ec11, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LinearRegression

# YOUR CODE HERE
pass
```

<!-- grade_id: cell-346d5411d0d8d874, points: 1, grade: true, solution: false -->
```python
from sklearn.metrics import mean_squared_error

assert isinstance(lr, LinearRegression)
assert np.isclose(y_preds.sum(), 264144945.99999997)
assert np.isclose(np.sqrt(mean_squared_error(y, y_preds)), 54087.066902509694)
```

## Part E. Build a random forest model.
Build a random forest model based on the Ames data. The features you will be using are: sqft_sum and Fireplaces (the number of fireplaces in the house).

Name the model *rf*.

Once the model is trained, predict the SalePrice for the same data you trained on. Name the predictions *y_preds*.

Hints:
* Use sklearn's RandomForestRegressor function to fit a model.

<!-- grade_id: cell-rfmodel, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import RandomForestRegressor
np.random.seed(42) # do not remove this line

# YOUR CODE HERE
pass
```

<!-- grade_id: cell-rfmodel-asserts, points: 1, grade: true, solution: false -->
```python
from sklearn.metrics import mean_squared_error

assert isinstance(rf, RandomForestRegressor)
assert np.isclose(np.sqrt(mean_squared_error(y, y_preds)), 28891.016368235392)
```

```python
x = sklearn.datasets.load_diabetes¶
```

```python

```